```python
%matplotlib inline
%precision 16
import numpy
import matplotlib.pyplot as plt
```

Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel $\rightarrow$ Restart) and then run all cells (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says YOUR CODE HERE or "YOUR ANSWER HERE", as well as your name and collaborators below:

# HW 4:  Numerical Differentiation and Quadrature

## Question 1 - Fundamental Theorem of Calculus

**(a)** (15) Write a function that computes the integral of the derivative 

$$I[f] = \int^x_0 f'(\hat{x}) d\hat{x}$$

and returns the error from the expected calculation.  Use a second order accurate centered difference scheme and then a left-hand rule to compute the absolute error.  In this case the left-hand quadrature is

$$Q[f] = f(x_i) \Delta x$$

The function should take in $x$ values and the function $f(x)$ and return a vector of the **errors** for each value $x$ given.

<!-- grade_id: A-1a, points: 0, grade: false, solution: true -->
```python
def int_diff_error(x, f):
    ### INSERT CODE HERE
    raise NotImplementedError("Replace this statement with your solution.")
    
    return error
```

<!-- grade_id: T-1a, points: 15, grade: true, solution: false -->
```python
f = lambda x: 1.0 / (1.0 + 100.0 * x**2)
x = numpy.linspace(0.0, 1.0, 500)
print numpy.max(int_diff_error(x, f))
numpy.testing.assert_allclose(int_diff_error(x, f), numpy.zeros(500), atol=1e-2)
print "Successful FTC evaluation."
```

**(b)** (15) Compute the same as above except reverse the order of the operations, i.e.

$$D[f] = \frac{\text{d}}{\text{d}x} \int^x_0 f(\hat{x}) d\hat{x}$$


<!-- grade_id: A-1b, points: 0, grade: false, solution: true -->
```python
def diff_int_error(x, f):
    ### INSERT CODE HERE
    raise NotImplementedError("Replace this statement with your solution.")
    
    return error
```

<!-- grade_id: T-1b, points: 15, grade: true, solution: false -->
```python
f = lambda x: 1.0 / (1.0 + 100.0 * x**2)
x = numpy.linspace(0.0, 1.0, 500)
print numpy.max(diff_int_error(x, f))
numpy.testing.assert_allclose(diff_int_error(x, f), numpy.zeros(500), atol=1e-2)
print "Successful FTC evaluation."
```

**(c)** (10) Plot the convergence rate for the number of partitions `N = [10,100,200,300,400,500]` on a `loglog` plot by computing the norm over the error given for each $N$ by using the `numpy.linalg.norm` command with `ord=2` vs. the $\Delta x$ used.  Theorize about what you observe regarding the order of convergence.

<!-- grade_id: A-1c1, points: 5, grade: true, solution: true -->


<!-- grade_id: A-1c2, points: 5, grade: true, solution: true -->
```python
### INSERT CODE HERE
raise NotImplementedError("Replace this statement with your solution.")
```

## Question 2 - Quadrature

Consider the scaled Fresnel integrals

$$I_c = \int^1_0 \frac{\cos x}{\sqrt{x}} dx~~~~\text{and}~~~~I_s = \int^1_0 \frac{\sin x}{\sqrt{x}} dx$$

which have values

$$I_c = \sqrt{2 \pi} \cdot C\left(\sqrt{\frac{2}{\pi}}\right) \approx 1.8090484758005438$$

$$I_s = \sqrt{2 \pi} \cdot S\left(\sqrt{\frac{2}{\pi}}\right) \approx 0.62053660344676231$$

where the functions $C(x)$ and $S(x)$ can be evaluated by `scipy.special.fresnel`.

**(a)** (10) Write a function that computes both integrals using the  trapezoidal rule with N partitions of equal length "ignoring" the singularity at $x=0$ by setting integrands to 0.

<!-- grade_id: A-2a, points: 0, grade: false, solution: true -->
```python
def trap_1(N):
    ### INSERT CODE HERE
    raise NotImplementedError("Replace this statement with your solution.")
    
    return Is, Ic
```

<!-- grade_id: T-2a, points: 10, grade: true, solution: false -->
```python
import scipy.special
Is, Ic = scipy.special.fresnel(numpy.sqrt(2.0 / numpy.pi))
Is *= numpy.sqrt(2.0 * numpy.pi)
Ic *= numpy.sqrt(2.0 * numpy.pi)
Is_hat, Ic_hat = trap_1(200)
error = numpy.abs(Is_hat - Is)
print "Error: %s" % error
numpy.testing.assert_allclose(error, 0.0, atol=1e-4)
print "Computed the integrals correctly using the trapezoidal rule only."    
```

**(b)** (10) A weighted Newton-Cotes quadrature rule is a modification of our version of Newton-Cotes quadrature except that we also multiply by a weighting function such that

$$\int^1_0 w(x) f(x) dx = \sum^{N}_{i=1} w_i f(x_i)$$

For the weight $x^{-1/2}$ we can show that 

$$\int^1_0 \frac{f(x)}{\sqrt{x}} dx \approx \frac{2}{3} (2 f(0) + f(1))$$

Combining this formula with the trapezoidal rule approach above, evaluate the integrals again by using the trapezoid rule except for the partition $[0,\Delta x]$ that involves the singularity use the above weighted Newton-Cotes rule.  Note that you will have to map the rule to the partition.

<!-- grade_id: A-2b, points: 0, grade: false, solution: true -->
```python
def trap_wnc(N):
    ### INSERT CODE HERE
    raise NotImplementedError("Replace this statement with your solution.")
    
    return Is, Ic
```

<!-- grade_id: T-2b, points: 10, grade: true, solution: false -->
```python
import scipy.special
Is, Ic = scipy.special.fresnel(numpy.sqrt(2.0 / numpy.pi))
Is *= numpy.sqrt(2.0 * numpy.pi)
Ic *= numpy.sqrt(2.0 * numpy.pi)
Is_hat, Ic_hat = trap_wnc(200)
error = numpy.abs(Is_hat - Is)
print "Error: %s" % error
numpy.testing.assert_allclose(error, 0.0, atol=1e-4)
print "Computed the integrals correctly using the trapezoidal and Newton-Cotes rules."    
```

**(c)** (10) Do a change of variables $x = t^2$ and evaluate each integral using the trapezoidal rule for the entire domain.

<!-- grade_id: A-2c, points: 0, grade: false, solution: true -->
```python
def trap_transformed(N):
    ### INSERT CODE HERE
    raise NotImplementedError("Replace this statement with your solution.")
    
    return Is, Ic
```

<!-- grade_id: T-2c, points: 10, grade: true, solution: false -->
```python
import scipy.special
Is, Ic = scipy.special.fresnel(numpy.sqrt(2.0 / numpy.pi))
Is *= numpy.sqrt(2.0 * numpy.pi)
Ic *= numpy.sqrt(2.0 * numpy.pi)
Is_hat, Ic_hat = trap_transformed(200)
error = numpy.abs(Is_hat - Is)
print "Error: %s" % error
numpy.testing.assert_allclose(error, 0.0, atol=1e-5)
print "Computed the integrals correctly using the trapezoidal rule."    
```

**(d)** (10) Do the same as in part (c) but use 3-point Gauss-Legendre quadrature.

<!-- grade_id: A-2d, points: 0, grade: false, solution: true -->
```python
def gauss_legendre_3(N):
    ### INSERT CODE HERE
    raise NotImplementedError("Replace this statement with your solution.")
    
    return Is, Ic
```

<!-- grade_id: T-2d, points: 10, grade: true, solution: false -->
```python
import scipy.special
Is, Ic = scipy.special.fresnel(numpy.sqrt(2.0 / numpy.pi))
Is *= numpy.sqrt(2.0 * numpy.pi)
Ic *= numpy.sqrt(2.0 * numpy.pi)
Is_hat, Ic_hat = gauss_legendre_3(10)
error = numpy.abs(Is_hat - Is)
print "Error: %s" % error
numpy.testing.assert_allclose(error, 0.0, atol=1e-10)
print "Computed the integrals correctly using the trapezoidal and Newton-Cotes rules."    
```

**(e)** (10) Check the order of accuracy for each of the methods above and plot these on a `loglog` plot with appropriate reference lines to indicate the order of each.

<!-- grade_id: A-2e, points: 10, grade: true, solution: true -->
```python
### INSERT CODE HERE
raise NotImplementedError("Replace this statement with your solution.")
```

**(f)** (10) Comment on the results of this question and the vast differences in convergence rates even between the two integrations (why is $I_c$ less accurate than $I_s$).

<!-- grade_id: A-2f, points: 10, grade: true, solution: true -->
